```python
%pylab inline
from numpy.testing import *
```

Output:
```
Populating the interactive namespace from numpy and matplotlib
```

Write a function that given the size $n$ generates the lower triangular matrix $L$, such that:
$$
L_{ij} = i+j+1\ \quad\ \mathrm{if}\ j\leq i\mathrm{, wih}\ i = 0 \ldots n-1
$$

*Encouragement:* try to write a *one-liner* that generates the matrix. To do that focus first on the $i$-th line of the matrix. Think about it as the composition of two lists `[...]+[...]`. The first one is composed by `i+j+1` where `j` ranges from `0` to `i+1`. The second one is a list of zeros ranging from `i+1` to `n`.

<!-- grade_id: lower_tri_0, points: 0, grade: false, solution: true -->
```python
def lower_tri_matrix(n):
    return matrix([append(arange(i+1,2*i+2), zeros(n-i-1)) for i in range(n)])
bb = lower_tri_matrix(4)
print(bb)
```

Output:
```
[[ 1.  0.  0.  0.]
 [ 2.  3.  0.  0.]
 [ 3.  4.  5.  0.]
 [ 4.  5.  6.  7.]]
```

<!-- grade_id: lower_tri_0_test, points: 2, grade: true, solution: false -->
```python
assert_equal(type(lower_tri_matrix(3)),numpy.matrixlib.defmatrix.matrix)
assert_equal(lower_tri_matrix(5),[[1,0,0,0,0],[2,3,0,0,0],[3,4,5,0,0],[4,5,6,7,0],[5,6,7,8,9]])
```

Write a function that solves a lower triangular linear system:
$$
L x = b
$$
Reminder:
$$
x_i = \frac{1}{L_{ii}}\left(b_i - \sum_{j=0}^{i-1} L_{ij}\, x_j\right)
$$

<!-- grade_id: lower_tri_1, points: 0, grade: false, solution: true -->
```python
def forward_substitutions(L,b):
    x = zeros(len(b))
    for i in range(len(b)):
        x[i] = (b[i] - einsum('ij,j', L, x)[i]) / L[i, i]
    return x
```

<!-- grade_id: lower_tri_1_test, points: 2, grade: true, solution: false -->
```python
L = lower_tri_matrix(5)
b = ones(5)
assert_almost_equal(forward_substitutions(L,b),solve(L,b),decimal=14)
```

Write a function that given the size $n$ generates the upper triangular matrix $U$, such that:
$$
U_{ij} = i+j+1\ \quad\ \mathrm{if}\ j\geq i\mathrm{, wih}\ i = 0 \ldots n-1
$$

*Encouragement:* try to write a *one-liner* that generates the matrix.

<!-- grade_id: upper_tri_0, points: 0, grade: false, solution: true -->
```python
def upper_tri_matrix(n):
    return matrix([append(zeros(i), arange(2*i+1,n + i + 1)) for i in range(n)])
```

<!-- grade_id: upper_tri_0_test, points: 2, grade: true, solution: false -->
```python
assert_equal(upper_tri_matrix(5),[[1,2,3,4,5],[0,3,4,5,6],[0,0,5,6,7],[0,0,0,7,8],[0,0,0,0,9]])
```

Write a function that solves a upper triangular linear system:
$$
U x = b
$$
Reminder:
$$
x_i = \frac{1}{U_{ii}}\left(b_i - \sum_{j=i}^{n-1} U_{ij}\, x_j\right)
$$

<!-- grade_id: upper_tri_1, points: 0, grade: false, solution: true -->
```python
def backward_substitutions(U,b):
    x = zeros(len(b))
    for i in reversed(range(len(b))):
        x[i] = (b[i] - einsum('ij,j', U, x)[i]) / U[i, i]
    return x
```

<!-- grade_id: upper_tri_1_tests, points: 2, grade: true, solution: false -->
```python
U = upper_tri_matrix(5)
b = ones(5)
assert_almost_equal(backward_substitutions(U,b),solve(U,b),decimal=14)
```

Write a function that performs the $LU$ factorisation of a given matrix $A$.

### Algorithm, Gaussian Elimination without pivoting, hints:
Set:
$$
U = A, \quad L = I.
$$
Choosing appropriate indexing:
$$
l_{ik} = \frac{u_{ik}}{u_{kk}}\\
u_{j,k:m} = u_{j,k:m} - l_{jk}\cdot u_{m,k:m}
$$


<!-- grade_id: lu_0, points: 0, grade: false, solution: true -->
```python
def lu_factorization(A):
    U = copy(A)
    N = A.shape[0]
    L = identity(N)
    for i in range(0, N):
        for j in range(i, N-1):
            L[j+1, i] = U[j+1, i] / U[i, i]
            U[j+1, i] = 0.
        
    #print(L)
        for k in range(i+1, N):
            for j in range(i+1, N):
                U[k, j] -= L[k,i] * U[i, j]
    return L, U
                
AAA = random.rand(5,5)
#print(AAA)
L, U = lu_factorization(AAA)
print(L, U)
```

Output:
```
[[ 1.          0.          0.          0.          0.        ]
 [ 6.07094824  1.          0.          0.          0.        ]
 [ 7.8678098   1.19389733  1.          0.          0.        ]
 [ 5.51348849  0.69913192 -1.34634757  1.          0.        ]
 [ 0.99859861 -2.1887442   8.52452094  0.3349631   1.        ]] [[ 0.09183257  0.0464986   0.20377476  0.12623684  0.83293456]
 [ 0.         -0.26073122 -0.65790923 -0.26693208 -4.09353669]
 [ 0.          0.         -0.16550727 -0.04897749 -1.34697117]
 [ 0.          0.          0.          0.13247981 -3.31718501]
 [ 0.          0.          0.          0.          3.04807565]]
```

<!-- grade_id: lu_0_tests, points: 2, grade: true, solution: false -->
```python
A = random.rand(5,5)
L,U = lu_factorization(A)
assert_almost_equal(amax(A - L.dot(U)),0,decimal=14)
```

Write a function to solve the linear system $A\,x=b$, via $LU$ factorization.

<!-- grade_id: lu_1, points: 0, grade: false, solution: true -->
```python
def lu_solve(A,b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: lu_1_tests, points: 2, grade: true, solution: false -->
```python
A = random.rand(5,5)
b = ones((5,))
assert_almost_equal(lu_solve(A,b),solve(A,b),decimal=14)
```