```python
%pylab inline
from numpy.testing import *
```

Write a function that given the size $n$ generates the lower triangular matrix $L$, such that:
$$
L_{ij} = i+j+1\ \quad\ \mathrm{if}\ j\leq i\mathrm{, wih}\ i = 0 \ldots n-1
$$

*Encouragement:* try to write a *one-liner* that generates the matrix. To do that focus first on the $i$-th line of the matrix. Think about it as the composition of two lists `[...]+[...]`. The first one is composed by `i+j+1` where `j` ranges from `0` to `i+1`. The second one is a list of zeros ranging from `i+1` to `n`.

<!-- grade_id: lower_tri_0, points: 0, grade: false, solution: true -->
```python
def lower_tri_matrix(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: lower_tri_0_test, points: 2, grade: true, solution: false -->
```python
assert_equal(type(lower_tri_matrix(3)),numpy.matrixlib.defmatrix.matrix)
assert_equal(lower_tri_matrix(5),[[1,0,0,0,0],[2,3,0,0,0],[3,4,5,0,0],[4,5,6,7,0],[5,6,7,8,9]])
```

Write a function that solves a lower triangular linear system:
$$
L x = b
$$
Reminder:
$$
x_i = \frac{1}{L_{ii}}\left(b_i - \sum_{j=0}^{i-1} L_{ij}\, x_j\right)
$$

<!-- grade_id: lower_tri_1, points: 0, grade: false, solution: true -->
```python
def forward_substitutions(L,b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: lower_tri_1_test, points: 2, grade: true, solution: false -->
```python
L = lower_tri_matrix(5)
b = ones(5)
assert_almost_equal(forward_substitutions(L,b),solve(L,b),decimal=14)
```

Write a function that given the size $n$ generates the upper triangular matrix $U$, such that:
$$
U_{ij} = i+j+1\ \quad\ \mathrm{if}\ j\geq i\mathrm{, wih}\ i = 0 \ldots n-1
$$

*Encouragement:* try to write a *one-liner* that generates the matrix.

<!-- grade_id: upper_tri_0, points: 0, grade: false, solution: true -->
```python
def upper_tri_matrix(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: upper_tri_0_test, points: 2, grade: true, solution: false -->
```python
assert_equal(upper_tri_matrix(5),[[1,2,3,4,5],[0,3,4,5,6],[0,0,5,6,7],[0,0,0,7,8],[0,0,0,0,9]])
```

Write a function that solves a upper triangular linear system:
$$
U x = b
$$
Reminder:
$$
x_i = \frac{1}{U_{ii}}\left(b_i - \sum_{j=i}^{n-1} U_{ij}\, x_j\right)
$$

<!-- grade_id: upper_tri_1, points: 0, grade: false, solution: true -->
```python
def backward_substitutions(U,b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: upper_tri_1_tests, points: 2, grade: true, solution: false -->
```python
U = upper_tri_matrix(5)
b = ones(5)
assert_almost_equal(backward_substitutions(U,b),solve(U,b),decimal=14)
```

Write a function that performs the $LU$ factorisation of a given matrix $A$.

### Algorithm, Gaussian Elimination without pivoting, hints:
Set:
$$
U = A, \quad L = I.
$$
Choosing appropriate indexing:
$$
l_{ik} = \frac{u_{jk}}{u_{kk}}\\
u_{j,k:m} = u_{j,k:m} - l_{jk}\cdot u_{m,k:m}
$$


<!-- grade_id: lu_0, points: 0, grade: false, solution: true -->
```python
def lu_factorization(A):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: lu_0_tests, points: 2, grade: true, solution: false -->
```python
A = random.rand(5,5)
L,U = lu_factorization(A)
assert_almost_equal(amax(A - L.dot(U)),0,decimal=14)
```

Write a function to solve the linear system $A\,x=b$, via $LU$ factorization.

<!-- grade_id: lu_1, points: 0, grade: false, solution: true -->
```python
def lu_solve(A,b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: lu_1_tests, points: 2, grade: true, solution: false -->
```python
A = random.rand(5,5)
b = ones((5,))
assert_almost_equal(lu_solve(A,b),solve(A,b),decimal=14)
```